<img src="https://brand.umich.edu/assets/brand/style-guide/logo-guidelines/U-M_Logo-Horizontal-Hex.png" alt="Drawing" style="width: 300px;" align="left"/><br>
    
## Week 1: Exploratory data analysis for learning analytics (20pts)

Welcome aspiring data scientists!

In this course, you will be working with a wide variety of data about students and their learning processes. In each week, you will be tasked with different questions that requires extracting insights from the data, in order to inform future teaching and learning practices. Here's an overview of the next four weeks:


| Stage    	| Task                                             	| Purpose                                                                                                               	|
|:--------	|:--------------------------------------------------	|:------------------------------------------------------------------------------------------------------------------------	|
| Week 1 	| Exploratory data analysis for learning analytics 	| What are some historical trends in students learning behavior and academic performance?                                	|
| Week 2 	| Building prediction models of student success    	| Which students are likely to dropout or fail? Can we build an early at-risk detection model for in time interventions? 	|
| Week 3 	| Investigating the biases of prediction models    	| Are the prediction models biased toward certain population of students? Are the data biased?                           	|
| Week 4 	| Building a learning analytics dashboard          	| How can we communicate learning analytics insights to a non-technical audience (instructors, students, managers)?      	|



**In week 1**, you should familiarize yourself with the datasets, which will be used throughout this course. We will use a public dataset from the Open University Learning Analytics (OULAD). The Open University UK is the largest online distance higher education institution in Europe with over 120,000 enrolled students. This dataset has been used in previous "hackathons" at the Learning Analytics Knowledge (LAK) conference, a premier venue for learning analytics research where many Michigan faculty members and researchers have participated in. 

**Overview of the dataset**

The dataset contains the information about 22 courses, 32,593 students, their assessment results, and logs of their interactions with the Virtual Learning Environment (e.g., Moodle) represented by daily summaries of student clicks (10,655,280 entries). 

**Reference**

Kuzilek, J., Hlosta, M., & Zdrahal, Z. (2017). Open university learning analytics dataset. Scientific data, 4, 170171. https://www.nature.com/articles/sdata2017171

# Open University Learning Analytics (OULAD) dataset

Kuzilek, J., Hlosta, M., & Zdrahal, Z. (2017). Open university learning analytics dataset. Scientific data, 4, 170171. https://www.nature.com/articles/sdata2017171
## Data scheme
![](https://media.springernature.com/lw685/springer-static/image/art%3A10.1038%2Fsdata.2017.171/MediaObjects/41597_2017_Article_BFsdata2017171_Fig2_HTML.jpg)
## Course timeline
![](https://media.springernature.com/lw685/springer-static/image/art%3A10.1038%2Fsdata.2017.171/MediaObjects/41597_2017_Article_BFsdata2017171_Fig1_HTML.jpg)
## Relational database
* A module is a course
* A presentation is a semester (e.g., 2019J - Fall 2019, 2019B = Winter 2019)
* vle = virtual learning enviroment
![](https://analyse.kmi.open.ac.uk/resources/images/model.png)

---

## Part A (4 points)
Write a function that returns a pandas dataframe that consists of passing rate and dropout rate for each course.   
The output data frame should have 4 columns: code_module, code_presentation, passrate, dropoutrate  
Note:
* The dropout rate should be relative to the total number of enrolled students
* The passing rate should be relative to the number students who completed the course (i.e. not Withdrawn)

<!-- grade_id: 1A_answer, points: 0, grade: false, solution: true -->
```python
"""The output data frame should have 4 columns: code_module, code_presentation, pass_rate, dropout_rate"""

import pandas as pd
import numpy as np

def answer_a():
    stdinfo = pd.read_csv('assets/studentInfo.csv')
    
    # YOUR CODE HERE
    raise NotImplementedError()
    
    return result
answer_a().head(5)
```

<!-- grade_id: 1A_solution, points: 4, grade: true, solution: false -->
```python
import pandas as pd
import numpy as np

# Check the number of courses
assert len(answer_a()) == 22

# Check the pass rate of AAA in 2013J
actual =  answer_a()[(answer_a()['code_module']=='AAA') &
                     (answer_a()['code_presentation']=='2013J')]['pass_rate'][0] 
desired = 0.860681
np.testing.assert_almost_equal(actual, desired, decimal=4, err_msg='', verbose=True)

```

```python
# let's visualize the courses by passrate and dropout rate
# Create a scatter plot
import seaborn as sns
sns.scatterplot(x="pass_rate", y="dropout_rate", data=answer_a(), hue="code_presentation")
```

---

## Part B (8 points)

As course instructor, we would be interested in tracking the level of student engagement on a weekly basis.  
Write a function that returns a pandas dataframe that consists of a weekly engagement metric, defined as the **average number of click per student per week** for each course.     
The output data frame should have 4 columns: **code_module, code_presentation, week, mean_click**.   
Use the following bin range to **convert date to week**. Hint: There are many ways to bin data, one of them is to use np.digitize and np.linspace. 


| date_bin 	| week 	|
|----------	|------	|
| ...      	| ...  	|
| [-14,-8] 	| -2   	|
| [-7,-1]  	| -1   	|
| [0,6]    	| 1    	|
| [7,13]   	| 2    	|
| [14,20]  	| 3    	|
| ...      	| ...  	|

<!-- grade_id: 1B_answer, points: 0, grade: false, solution: true -->
```python
"""The output data frame should have 4 columns: code_module, code_presentation, week, mean_click"""
import pandas as pd
import numpy as np

def answer_b():
    result = None
    # YOUR CODE HERE
    raise NotImplementedError()
    return result
answer_b()
```

<!-- grade_id: 1B_solution, points: 8, grade: true, solution: false -->
```python
# Autograder tests

result_b = answer_b()

# Check the length of output
assert len(result_b) == 879

# Check the mean_click of AAA in 2013J week 1
actual =  result_b[(result_b['code_module']=='AAA') &
                     (result_b['code_presentation']=='2013J') &
                     (result_b['week']==1)]['mean_click'].values
desired = 100.771513
np.testing.assert_almost_equal(actual, desired, decimal=2)

# Check the week
if 0 in result_b['week'].unique():
    raise AssertionError("There should be no week 0 in the dataset, check the instruction") 

```

```python
# Let's plot the result
import seaborn as sns
sns.relplot(data=answer_b(), x="week", y="mean_click", col="code_module",hue="code_presentation",
            kind="line", col_wrap=4)
```

---

## Part C (8 points)

As course instructors, we would be interested in tracking the submission rate and the pass rate of each assignment.  
Write a function that returns a pandas dataframe that consists of 6 columns: **code_module, code_presentation, id_assessment, assessment_type, submission_rate, pass_rate**.

**Submission_rate** is defined as the **number of students who submitted the assignment** divided by the **current number of enrollment**. The **current number of enrollment** changes over time as students **withdraw/unregister**.

For example, course A in week 1 has 100 enrolled students, 90 submitted their week 1 assignment, submission rate week 1 = 90%
In week 2, 5 students dropped out, so the enrollment number is now 95. Out of 95 students, 85 submitted week 2 assignment, submission rate week 2 = 85/95 = 89.5%

**Passing_rate** is defined as the **number of students who submitted the assignment AND achieve score >50** divided by the **number of students who submitted the assignment**.

Note: 
* Replace the missing values in the 'date' column in assessments.csv with 250
* Delete row where id_assessment == 40088 & 40087

<!-- grade_id: 1C_answer, points: 0, grade: false, solution: true -->
```python
"""The output data frame should have 6 columns: code_module, code_presentation, id_assessment, assessment_type, submission_rate, pass_rate"""
import pandas as pd
import numpy as np

def answer_c():
    # Import assessments.csv
    assessments = pd.read_csv('assets/assessments.csv')
    # Data cleaning
    assessments["date"] = assessments["date"].fillna(250)
    assessments = assessments[(assessments['id_assessment']!=40088) &
                              (assessments['id_assessment']!=40087)]
    # Import student registration data
    stdreg = pd.read_csv('assets/studentRegistration.csv')
    # Import student grades
    stdgrade = pd.read_csv('assets/studentAssessment.csv')
    
    # YOUR CODE HERE
    raise NotImplementedError()
    return result
answer_c()
```

<!-- grade_id: 1C_solution, points: 8, grade: true, solution: false -->
```python
import pandas as pd
import numpy as np

result_c = answer_c()

# Check the length of output
assert len(result_c) == 204

# Check the submission rate of AAA in 2013J id_assessment==1756
actual =  result_c[(result_c['code_module']=='AAA') &
                     (result_c['code_presentation']=='2013J') &
                     (result_c['id_assessment']==1756)]['submission_rate'].values
desired = 0.908537
np.testing.assert_almost_equal(actual, desired, decimal=2)

```